```python
# Initialize Otter
import otter
grader = otter.Notebook("ps1.ipynb")
```

# STATS 507 W22
## Problem Set 1
We'll be using the `math` module a few times:

```python
import math
```

All functions will be tested by visible as well as hidden tests. The maximum amount of time any function is allowed to run is 10 seconds.

### Question 1: Euclid's algorithm 
Euclid's algorithm (<https://en.wikipedia.org/wiki/Euclidean_algorithm>)
is a method for finding the greatest common divisor (GCD) of two
numbers. Recall that the GCD of two numbers $m$ and $n$ is the largest
number that divides both $m$ and $n$.

**1(a)** (4 pts.) The Wikipedia page above includes several pseudocode implementations
    of Euclid's algorithm. Choose one of these, and use it to implement
    a function `gcd(m,n)`, which takes two integers as its arguments and
    returns their GCD. You may assume that both inputs are integers, so
    there is no need to include any error checking in your function. 
    
   **Note**: There are many existing library functions, such as `math.gcd()`, available for calculating this. Do not use them: your job is to implement this function by hand.

```python
def gcd(m, n):
    "Returns the greatest common divisor of integers m and n."
    if n == 0:
        return m
    else:
        return max(gcd(n, m % n),-gcd(n, m % n))
    ...
```

```python
grader.check("q1a")
```

Output:
```
q1a results: All test cases passed!
```

<!-- BEGIN QUESTION -->

**1(b)** (1 pt.) What does your function do if one or both of its arguments are negative? Does this behavior make sense?

For Euclidian division it does not make sense we should only allow non-negative values. But with the code above, we can use negative values in the same way.

<!-- END QUESTION -->

### Question 2: Approximating $e$
The base of the natural logarithm, $e$, is typically defined as the
infinite sum 
$$
e = \sum_{k=0}^\infty \frac{1}{k!}
= 1 + 1 + \frac{1}{2} + \frac{1}{6} + \frac{1}{24} + \dots,$$ where $k!$ denotes the factorial of $k$,
$$k! = k \cdot (k-1) \cdot (k-2) \cdot \dots \cdot 3 \cdot 2 \cdot 1,$$ where we define $0!=1$ by convention. For more on Euler's number, see
<https://en.wikipedia.org/wiki/E_(mathematical_constant)>. In this
problem, we will explore different approaches to approximating this
number.


**2(a)** (2 pts.) An early characterization of Euler's number, due to Jacob Bernoulli,
    was as the limit
    $$e = \lim_{x\to\infty} \left(1 + \frac{1}{x}\right)^x.$$ Define a function called `e_limit(n)` that takes as an argument an integer $n$, and returns a `float` that approximates $e$ by taking $x=n$ in the above equation. You may assume that the input to your function will be a positive integer (of course, your function will still run just fine if $n$ is, say, a positive float, but we will only use integer values in what follows).

```python
def e_limit(n):
    "Approximates Euler's number using the limit definition."
    e = (1+1/n)**n
    return e
    ...
```

```python
grader.check("q2a")
```

Output:
```
q2a results: All test cases passed!
```

<!-- BEGIN QUESTION -->

**2(b)** (1 pt.) What happens when you call `e_lim(n)` for really huge values of `n` (say,
    $10^{16}$ or $10^{18}$)? Why does this happen? **Hint:** the answer
    has to do with floating point arithmetic. Think about how $1/n$ is
    represented on your computer when $n$ is big.

The function will return 1.0 because the computer will represent 1/n as 0 as it currently doesn't have memory to store 1/(10^18) number and considers it a zero.
Honestly I believe that better implementation of this function is to use a while loop like this:
```
def e_limit(n):
    "Approximates Euler's number using the limit definition."
    while n < 1000000:
        e = (1+1/n)**n
        n = n+1
    return e
```
and make an exception for larger numbers that n will increase by larger increments.

<!-- END QUESTION -->

**2(c)** (2 pts.) Define a function called `e_series` that takes a single non-negative integer
    argument `n`, and returns an approximation to $e$ based on the first
    `n` terms of the infinite sum shown above. Your function should return a float. You may again assume that the input will be a non-negative integer, so you do not need to include error checking in your function. As an example, `e_series(4)` should return the sum of the first four terms the series, i.e. $1 + 1 + 1/2 + 1/6 \approx 2.667$. 

Notes:

- The sum the equation starts counting with $k=0$ (i.e., it is "$0$-indexed"), while our function starts counting with $n=1$ (i.e., it is "$1$-indexed"). `e_series(1)` should use *one* term, so that `e_series(1) == 1`. Similarly, `e_series(0)` should return $0$, since by convention an empty sum is equal to zero.
- You may use the function `math.factorial(k)` to compute $k!$, but I
    recommend, for the sake of practice, implementing the factorial
    function yourself, as it is a nice example of a problem that is
    easily solved with recursion.

```python
def e_series(n):
    "Approximates Euler's number using the series definition."
    e = 0
    for i in range(n):
        e = e + 1/math.factorial(i)
    return e
    ...
```

```python
grader.check("q2c")
```

Output:
```
q2c results: All test cases passed!
```

**2(d)** (2 pts.) Define a function called `error_table(N)` takes a single positive
    integer `N` as an argument and returns a data structure containing the _relative_ error between the
    approximations given by, respectively, `e_limit` and `e_series`. `error_table(N)` should return a
    list of length $N$. The $k$-th entry of the list, `error_table(N)[k]`, should be a `dict` with two keys, `err_limit` and `err_series`, giving the relative errors for each approach when $n=k+1$:
```
>>> error_table(1)
[{'err_limit': 0.26424111765711533, 'err_series': 0.6321205588285577}]
```
(The first entry of the list is shown.)
**Note**: the true value of $e$, accurate to machine precision, can be obtained as `math.e`.

```python
def error_table(N):
    "Table of relative errors for e_limit(k) and e_series(k) for k=1,...,N."
    err_tab = []
    for i in range(N):
        err_limit = (math.e-e_limit(i+1))/math.e
        err_series = (math.e-e_series(i+1))/math.e
        err_tab.append({'err_limit': err_limit, 'err_series': err_series})
    return err_tab
    ...
```

```python
grader.check("q2d")
```

Output:
```
q2d results: All test cases passed!
```

<!-- BEGIN QUESTION -->

**2(e)** (1 pt.) Judging from the table, which approximation method is better? Why do you think this is?

err_series approximation method is better as it has smaller relative error. It happens because it is a numerical/iterative approximation rather than an analytical/empirical appx, in which we include more steps to ensure we are as close to the number as possible.

<!-- END QUESTION -->

### Question 3: Testing properties of an integer
In this problem, you'll get a bit more practice working with
conditionals, and a first exposure to the kind of thinking that is
required in a typical "coding interview" question. A positive integer
$n$ is a *power of $2$* if $n = 2^p$ for some integer $p \ge 0$.

**3(a)** (2 pts.) Using only the division and modulus (`%`)
    operators, write a function `is_pow_2(n)` that takes a positive integer $n$ as its only
    argument and returns a `bool` indicating whether or not the input
    is a power of 2. You may assume that the input is a positive
    integer. **Hint:** the simplest solution to this problem makes
    use of recursion, though recursion is not necessary.

```python
def is_pow_2(n):
    "Returns True if n = 2 ** p for some integer p >= 0."
    while n != 1: 
        if n % 2 == 0: 
            n = n / 2
        else:
            return False
    return True
    ...
```

```python
grader.check("q3a")
```

Output:
```
q3a results: All test cases passed!
```

**3(b)** (2 pts.)
Generalize your previous solution to a function `is_pow_k` that takes two
    positive integers $n$ and $k$ as its arguments, in that order, and
    returns `True` if `n` is a power of `k`
    and `False` otherwise.

```python
def is_pow_k(n, k):
    "Returns True if n == k ** p for some integer p >= 0."
    if k == 1:
        return n
    while n != 1 or 0: 
        if n % k == 0:
            n = n / k
        else:
            return False
    return True
    ...

is_pow_k(1,1)
```

Output:
```
1
```

```python
grader.check("q3b")
```

Output:
```
q3b results: All test cases passed!
```

### Question 4: Root finding and minimization
Recall that a *root* $r$ of a function $f:\mathbb{R}\to\mathbb{R}$ satisfies $f(r)=0$. The [bisection method](https://en.wikipedia.org/wiki/Bisection_method) is a generic algorithm for finding the root of a continuous, real-valued function. 

**4(a)** (4 pts.) Write a function `find_root(f, a, b)` which uses the bisection method to locate a root of the function $f$ on the interval $[a,b]$. Here, $f$ is a Python function which accepts a single float, and returns a float (i.e., a real-valued function of one variable.) You may assume that $f$ is continuous and that $f(a)$ and $f(b)$ have different signs, that is, `sgn(f(a)) != sgn(f(b))` where `sgn(x)` is the sign function:

```python
def sgn(x):
    if x < 0:
        return -1
    if x > 0:
        return 1
    return 0
```

(**Note**: due to finite precision issues mentioned above, your algorithm may not find an exact root of $f$, i.e. a number such that `f(r) == 0.0`. We will say that $r$ is a root if $|f(r)|<10^{-16}$.)

```python
def find_root(f, a, b):
    "Find a root of f(x) in the interval [a, b], assuming that sgn(f(a)) != sgn(f(b))."
    n = 1
    while n <= 100:
        c = (a + b)/2
        if f(c) == 0 or (b - a)/2 < 1e-16:
            break
        if sgn(f(c)) == sgn(f(a)):
            a = c
        else:
            b = c
        n = n+1
    return c
    ...
```

### grader.check("q4a")

<!-- BEGIN QUESTION -->

**4(b)** (2 pts.) Suppose that $f:[a,b]\to\mathbb{R}$ is a strictly convex and continuously differentiable function. Let $x^*\in [a,b]$ be such that $f(x^*) \le f(x)$ for all $x\in[a,b]$. Explain using calculus how to use the root finding algorithm you implemented above to find $x^*$.

I can firstly derive the function and then find its square root, which will return the minimum of the original function. Minimum of the function will always satisfy the above equation.

<!-- END QUESTION -->

**4(c)** (4 pts.) Write a function `minimize(f_prime, a, b)` that minimizes a function $f$ satisfying the assumptions of the previous question. Here `f_prime` is a function that returns the derivative of $f$ evaluated at the point $x$.

```python
def minimize(f_prime, a, b):
    """
    Minimizes the convex function f over the interval [a,b].
    
    Args:
        f_prime: a function representing the (continuous) derivative of f.
        a, b: interval on which f is defined.
        
    Returns:
        The point x at which f is minimized.
    """
    min = find_root(f_prime, a, b)
    return min
    ...
```

```python
grader.check("q4c")
```

Output:
```
q4c results: All test cases passed!
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Output:
```
q1a results: All test cases passed!

q2a results: All test cases passed!

q2c results: All test cases passed!

q2d results: All test cases passed!

q3a results: All test cases passed!

q3b results: All test cases passed!

q4a results: All test cases passed!

q4c results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Upload this .zip file to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Output:
```
<IPython.core.display.HTML object>
```

Tests for q1a (points: 4)
```
>>> gcd(0, 0) == 0
True
```
```
>>> for m in range(1, 10):
...     assert gcd(0, m) == m
```
```
>>> gcd(1, 1) == 1
True
```
```
>>> for m in range(-10, 10):
...     assert gcd(m, 1) == 1
...     assert gcd(m, 2) == 2 - (m % 2)
```


Tests for q2a (points: 2)
```
>>> e_limit(1) == 2.
True
```
```
>>> e_limit(2) == 2.25
True
```


Tests for q2c (points: 2)
```
>>> e_series(0) == 0
True
```
```
>>> e_series(1) == 1.
True
```


Tests for q2d (points: 2)
```
>>> for k in range(1, 5, 10):
...     assert len(error_table(k)) == k
```
```
>>> assert abs(error_table(1)[0]['err_limit'] - 0.26424111765711533) < 1e-6
>>> assert abs(error_table(1)[0]['err_series'] - 0.6321205588285577) < 1e-6
```


Tests for q3a (points: 2)
```
>>> is_pow_2(2) == True
True
```
```
>>> is_pow_2(1) == True
True
```


Tests for q3b (points: 2)
```
>>> for k in range(1, 20):
...     assert is_pow_k(k, k)
```


Tests for q4a (points: 4)
```
>>> f = lambda x: x ** 3
>>> r = find_root(f, -4, 3)
>>> assert abs(f(r)) < 1e-16
```
```
>>> r = find_root(math.cos, 0, 3.14)
>>> assert abs(math.cos(r)) < 1e-16
```


Tests for q4c (points: 4)
```
>>> f_prime = lambda x: 2 * x  # derivative of f(x) = x^2
>>> assert abs(minimize(f_prime, -10, 10) - 0) < 1e-8
```
```
>>> f_prime = lambda x: -math.cos(x)  # derivative of f(x) = -sin(x)
>>> assert abs(minimize(f_prime, 0, 3.14) - math.pi/2) < 1e-8
```
