```python
# Initialize Otter
import otter
grader = otter.Notebook("ps3.ipynb")
```

# STATS 507 W22
## Problem Set 3
All functions will be tested by visible as well as hidden tests. The maximum amount of time any function is allowed to run is 10 seconds.

You may use any module in the [standard library](https://docs.python.org/3/library/) to solve these problems. You may not (yet) use other modules (scipy, numpy, pandas, etc.) that require installation. Here are a few modules that may prove particularly useful:

```python
import collections
import string
import itertools
```

### Question 1: Infinite sequences
For each of the problems below, write a function which generates the given infinite sequence. We should be able to use your generators to access any entry of the sequence no matter how deep. 

**1(a)** (2 pts) The prime numbers are

$$2, 3, 5, 7, 11, 13, 17, \dots$$

Give a generator for the primes. (Note: many algorithms exist for this problem. Yours should be efficient enough that we can use it to generate reasonably large prime numbers.)

```python
def divisors(i):
    return [j for j in range(1,i+1) if j <= i]

def is_prime(i):
    return len([j for j in divisors(i) if divisors(i)[-1] % j == 0]) == 2

def primes():
    i = 0
    limit = 0
    while True:
        if limit < 200:
            if is_prime(i):
                yield i
            i += 1
            limit += 1
        else:
            break
    ...
```

```python
grader.check("q1a")
```

Output:
```
q1a results: All test cases passed!
```

**1(b)** (3 pts) The *ruler sequence* is

$$1, 2, 2, 3, 3, 4, 4, 4, 5, 5, 5, 6, 6, 6, 7, 7, 7, 7\dots$$

The first number (1) appears once; the next two numbers (2 and 3) appear twice, the next three numbers appear three times, etc.

```python
def ruler():
    rep_control = 1
    num = 1
    limit = 0
    while True:
        if limit < 200:
            for x in range(rep_control):
                for y in range(rep_control):
                    yield num
                num += 1
            rep_control += 1
        limit += 1
    ...
```

```python
grader.check("q1b")
```

Output:
```
q1b results: All test cases passed!
```

**1(c)** (3 pts) The look-and-say sequence  

$$1, 11, 21, 1211, 111221, 312211, 13112221, 1113213211, \dots$$

is generated as follows:

1. The first entry is one.
2. The second entry is generated by reading the first entry aloud: "one one"
3. The third entry is generated by reading the second entry aloud: "two ones"
4. The fourth entry is generated by reading the third entry aloud: "one two and one one"
5. The fifth entry is therefore "one one, one two, and two ones"

... and so forth. Note that each entry of the sequence should be a `str` object.


```python
def look_say():
    orig_num = 1
    skip = 0
    limit = 0
    while True:
        if limit < 200:
            if skip != 0:
                orig_num = str(orig_num)
                prev = orig_num[0]
                count = 0
                next_t = ""
                for t in orig_num:
                    if t == prev:
                        count += 1
                    else:
                        next_t += str(count) + prev
                        count = 1
                    prev = t
                next_t += str(count) + prev
                start = int(orig_num)
                yield next_t
                orig_num = next_t
            else:
                yield str(1)
            skip = 1
        limit += 1
    ...
```

```python
grader.check("q1c")
```

Output:
```
q1c results: All test cases passed!
```

**1(d)** (3 pts) July 2022 is a special month: it contains 5 Fridays, 5 Saturdays, and 5 Sundays. Write a generator that, starting with July, 2022, generates each successive special month (according to this definition) in order. The format of each generated entry should be a string of the form `"<month> <year>"`, for example `"July 2022"`, `"December 2845"`, etc. *Hint*: you may find the `calendar` module useful.

```python
import calendar
import datetime
def special_months():
    Friday = calendar.FRIDAY
    Saturday = calendar.SATURDAY
    Sunday = calendar.SUNDAY
    year = 2022
    month = 7
    limit = 0
    while True:
        if limit < 200:
            matrix = calendar.monthcalendar(year,month)
            Fridays = (sum(1 for x in matrix if x[Friday] != 0))
            Saturdays = (sum(1 for x in matrix if x[Saturday] != 0))
            Sundays = (sum(1 for x in matrix if x[Sunday] != 0))

            if Fridays == 5 and Saturdays == 5 and Sundays == 5:
                datetime_object = datetime.datetime.strptime(str(month), "%m")
                yield datetime_object.strftime("%B")+ " " + str(year)
            if month == 12:
                month = 1
                year += 1
            else:
                month += 1
            limit += 1
        else:
            break
    ...
```

Output:
```
'\nimport calendar\nimport datetime\ndef special_months():\n    Friday = calendar.FRIDAY\n    Saturday = calendar.SATURDAY\n    Sunday = calendar.SUNDAY\n    year = 2022\n    month = 7\n    limit = 0\n    while True:\n        if limit < 200:\n            matrix = calendar.monthcalendar(year,month)\n            Fridays = (sum(1 for x in matrix if x[Friday] != 0))\n            Saturdays = (sum(1 for x in matrix if x[Saturday] != 0))\n            Sundays = (sum(1 for x in matrix if x[Sunday] != 0))\n\n            if Fridays == 5 and Saturdays == 5 and Sundays == 5:\n                datetime_object = datetime.datetime.strptime(str(month), "%m")\n                yield datetime_object.strftime("%B")+ " " + str(year)\n            if month == 12:\n                month = 1\n                year += 1\n            else:\n                month += 1\n            limit += 1\n        else:\n            break\n    ...\n'
```

```python
grader.check("q1d")
```

Output:
```
q1d results: All test cases passed!
```

**1(e)** (3 pts) A palindromic date is a date that reads the same forwards and backwards when written in standard `M-D-YY` format. This year, the palindromic dates are 2-2-22 and 2-20-22 through 2-28-22. (Dates with single-digit years are written with no leading zero, e.g. 1-1-1 for January 1, 2001). Starting with today, write a generator that returns all future palindromic dates:

```
>>> g = palindromic_dates()
>>> next(g)
"February 2, 2022"
>>> next(g)
"February 20, 2022"
...
>>> next(g)
"February 28, 2022"
>>> next(g)
"March 2, 2023"
```

```python
import calendar
import datetime
def palindromic_dates():
    Friday = calendar.FRIDAY
    Saturday = calendar.SATURDAY
    Sunday = calendar.SUNDAY
    year = 2022
    month = 1
    date = 1
    limit = 0
    while True:
        if limit < 500:
            matrix = calendar.monthcalendar(year,month)
            for week in matrix:
                for day in week:
                    if day != '0':
                        date = str(month) + str(day) + str(year)[-2:]
                    if date == str(date)[::-1]:
                        datetime_object = datetime.datetime.strptime(str(month), "%m")
                        yield datetime_object.strftime("%B")+ " " + str(day) + ", " + str(year)
            if month == 12:
                month = 1
                year += 1
            else:
                month += 1
            limit += 1
        else:
            break
    ...
```

```python
grader.check("q1e")
```

Output:
```
q1e results: All test cases passed!
```

### Question 2: Simple ciphers

A *cipher* is an algorithm for encrypting or decrypting a text message, called *plaintext*, into *ciphertext*. In this exercise, we will assume that all plaintext to be encrypted or decrypted are composed of the lower-case Roman alphabet, without any punctuation or whitespace. Examples of such messages could be

```attackatdawn```

or 

```iamajellyfilleddonut```.

**2(a)** (4 pts) One of the oldest known ciphers is the [Caesar cipher](https://en.wikipedia.org/wiki/Caesar_cipher) or shift cipher, attributed to the Roman emperor. <img src="https://upload.wikimedia.org/wikipedia/commons/b/b4/Bust_of_Julius_Caesar_from_History_of_the_World_%281902%29.png" width=100 style="float: right; margin: 0 0 10px 10px;" /> The cipher works by shifting all letters of the alphabet by a pre-specified integer $k$. For example, if $k=+1$ then the plaintext

```thequickbrownfoxjumpsoverthelazydog```

encrypts to the ciphertext

```uifrvjdlcspxogpykvnqtpwfsuifmbazeph```

(Notice that in this example the letter `z` in `lazy` wrapped around to become an `a`.) 

Write two functions, `enc_caesar(s, k)` and `dec_caesar(s, k)` that respectively encrypt and decrypt the string `s` using the Caesar cipher, based on the integer key `k`.

```python
def enc_caesar(s: str, k: int) -> str:
    '''
    Encrypt the message s using Ceasar cipher with key k.
    
    >>> enc_caeasar('thequickbrownfoxjumpsoverthelazydog', 1)
    'uifrvjdlcspxogpykvnqtpwfsuifmbazeph'
    '''
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    c = list(s)
    for i in range(len(c)):
        index = alphabet.find(c[i])
        new_index = index + k
        if new_index > len(alphabet)-1:
            new_index = new_index - len(alphabet)
        c[i] = c[i].replace(c[i],alphabet[new_index])
    
    cipher = "".join(c)
    return cipher
    ...
    
def dec_caesar(s: str, k: int) -> str:
    '''
    Decrypt the message s using Ceasar cypher with key k.
    
    >>> dec_caeasar('uifrvjdlcspxogpykvnqtpwfsuifmbazeph', 1)
    'thequickbrownfoxjumpsoverthelazydog'
    '''
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    c = list(s)
    for i in range(len(c)):
        index = alphabet.find(c[i])
        new_index = index - k
        if new_index > len(alphabet)-1:
            new_index = new_index - len(alphabet)
        c[i] = c[i].replace(c[i],alphabet[new_index])
    
    decipher = "".join(c)
    return decipher
    ...
```

```python
grader.check("q2a")
```

Output:
```
q2a results: All test cases passed!
```

<img src="https://upload.wikimedia.org/wikipedia/commons/9/9a/Vigenère_square_shading.svg" style="float: right; margin: 0 0 10px 10px" width=300 />


**2(b)** (4 pts) The Caesar ciphers are easily broken since the key can either be guessed, or is often widely available. A much more secure cipher is the [Vigenère cipher](https://en.wikipedia.org/wiki/Vigenère_cipher), which is keyed using a pre-specified word. The cipher works using a [tabula recta](https://en.wikipedia.org/wiki/Tabula_recta), a picture of which is shown to the right. 

Given a message `s` to be encrypted, the key `k` is repeated until it has the same length as `s`. Then each entry of the encrypted message `e` is obtained by looking up the corresponding **row** in `k`, and the correspending  **column** in `s`, in the tabula recta.

For example, suppose the message is `hello`, and the key is `sun`. We repeat the key until it has the same length as hello: `sunsu`. Then we encrypt the message by looking up the each entry in the table: `(s,h) (u,e) (n,l) (s,l) (u,o)`. The resulting string is `zyydi`.

```python
def enc_vignere(s: str, k: str) -> str:
    "Encrypt s using Vignere cipher with key k"
    c = list(s)
    l = list(k)
    i2=0
    for i in range(len(c)):
        if i > len(l)-1:
            i2 = i - len(l)
            if i2 > len(l)-1:
                i2 = i2 - len(l)
        c[i] = c[i].replace(c[i], l[i2])
        i2 +=1
        
    c = "".join(c)
    
    tuples = []
    for q, w in zip(s,c):
        tuples.append((q,w))
    
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    
    num_tuples = []
    for x in tuples:
        y = list(x)
        y[0] = alphabet.find(y[0])
        y[1] = alphabet.find(y[1])
        num_tuples.append(sum(tuple(y)))
    
    cipher = []
    for i in num_tuples:
        if i > len(alphabet)-1:
            i = i - len(alphabet)
        cipher.append(alphabet[i])
    
    cipher = "".join(cipher)
    return cipher
    ...

def dec_vignere(s: str, k: str) -> str:
    "Decrypt s which was encrypted by Vignere cipher using key k"
    c = list(s)
    l = list(k)
    i2=0
    for i in range(len(c)):
        if i > len(l)-1:
            i2 = i - len(l)
            if i2 > len(l)-1:
                i2 = i2 - len(l)
        c[i] = c[i].replace(c[i], l[i2])
        i2 +=1
        
    c = "".join(c)
    
    tuples = []
    for q, w in zip(s,c):
        tuples.append((q,w))
    
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    
    num_tuples = []
    for x in tuples:
        y = list(x)
        y[0] = alphabet.find(y[0])
        y[1] = alphabet.find(y[1])
        num_tuples.append(y[0]-y[1])
    
    cipher = []
    for i in num_tuples:
        if i > len(alphabet)-1:
            i = i - len(alphabet)
        cipher.append(alphabet[i])
    
    cipher = "".join(cipher)
    return cipher
    ...
```

```python
grader.check("q2b")
```

Output:
```
q2b results: All test cases passed!
```

<!-- BEGIN QUESTION -->

**2(c)** Suppose you are an attacker trying to decrypt a message by randomly guessing the key. What is your chance of guessing the key correctly on the first try if the message is encrypted using the Caesar cipher? What about for the Vignére cipher?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**2(d)** (3 pts) A [one-time pad](https://en.wikipedia.org/wiki/One-time_pad) (OTP) is a special type of Vignére cipher where the key is a randomly generated string equal in length to the message that is to be encoded. It can be mathematically proved that an OTP is crytographically perfect, in the sense that the enciphered message provides no information about the original message. Convince yourself of this fact by performing the following experiment:

1. Pick a message of your choosing to be encrypted. It should be not too short or too long (~20 characters).
2. Repeat a large number of times:
    - Generate a random string of the same length as the message. (Hint: the `random.choices` method may be useful here.)
    - Encrypt the message using this string as the key.
    - Save the resulting encrypted message for later analysis.

Now consider the first letter of each of the encrypted messages that you generated in step 2. This is a random variable that has a particular distribution, from which we have drawn a large number of samples. Estimate the distribution of this random variable using a histogram.  (You may find the counting functions you wrote for PS1 useful here.) What if I consider the joint distribution of the first two letters (or any pair of letters more generally)---what is their joint distribution? Suppose you are an attacker who wants to decrypt the encrypted message: do the results of these simulations suggest that this is feasible? Why or why not? 

From the printed histogram we can see that all any first letter occurs in the large sample of encrypted messages about as many times as any other letter. Therefore I estimate the distribution of the first letters to have Uniform distribution in which all outcomes are equally likely. The same goes for two letters of the encrypted words.
We can indeed prove that OTP is cryptographically perfect, because an attacker doesn't have any pattern which he could use for decrypting the encrypted message without knowing the relevant key word.

```python
msg = "Thishomeworkishard"
import random
import string

encryption = []
for i in range(30000):
    key = ''.join(random.choices(string.ascii_lowercase, k=len(msg)))
    cipher = enc_vignere(msg,key)
    encryption.append(cipher)

first_letters = []
for i in encryption:
    first_letters.append(i[0])
    
from collections import Counter
Counter(first_letters)

two_first_letters = []
for i in encryption:
    two_first_letters.append(i[0:2])
    
Counter(two_first_letters)
```

Output:
```
Counter({'qv': 42,
         'lr': 46,
         'pt': 40,
         'yu': 58,
         'lq': 52,
         'nr': 48,
         'ev': 46,
         'ov': 54,
         'yn': 39,
         'tz': 50,
         'fq': 43,
         'mf': 48,
         'kx': 50,
         'mn': 43,
         'ka': 32,
         'zb': 55,
         'fy': 57,
         'rk': 51,
         'ez': 42,
         'ct': 53,
         'th': 44,
         'jh': 57,
         'ye': 51,
         'ng': 35,
         'vg': 43,
         'pa': 53,
         'hl': 31,
         'vs': 34,
         'vl': 49,
         'xt': 43,
         'hu': 35,
         'qk': 40,
         'qf': 41,
         'sw': 41,
         'of': 47,
         'to': 37,
         'yd': 55,
         'jt': 55,
         'xb': 44,
         'dg': 40,
         'ua': 48,
         'rn': 41,
         'yf': 43,
         'dd': 34,
         'sm': 51,
         'wm': 37,
         'vb': 42,
         'ci': 50,
         'zy': 54,
         'nj': 43,
         'st': 45,
         'fg': 53,
         'me
```

<!-- END QUESTION -->

### Question 3: Book ciphers
<img src="https://upload.wikimedia.org/wikipedia/commons/e/e8/King-James-Version-Bible-first-edition-title-page-1611.png" width=200 style="float: right; margin: 0 0 10px 10px" />

A [bible cipher](https://en.wikipedia.org/wiki/Book_cipher) uses the King James Bible (or some other widely available text) as the key. The cipher works by replacing each word in the message with a reference to a particular location in the Bible where that word occurs. 

The bible is organized hierarchically into *books*, *chapters*, and *verses*. For example, the first sentence in the King James Bible is 

>In the beginning God created the heaven and the earth. (Genesis 1:1)

This is the first sentence of the Book of Genesis, chapter 1, verse 1. In a bible cipher, an occurrence the word "beginning" in the plaintext message could be replaced by the tuple `(0,0,0,2)`, in reference to the third word of the first book, first chapter, and first verse of the Bible.

**3(a)** (3 pts) The file `kjv.txt` contains complete text of the King James Bible. 

Parse this file into a data structure `bible` 
such that calling `bible[book][chapter][verse][i]` returns the `i`th word of the corresponding book, chapter and verse. Here `book`, `chapter`, and `verse`, and `i` are integers representing the (0-indexed) position of corresponding word. You should convert all words to lowercase and remove anything that is not an alphabetic character or a space:

```
>>> bible[0][0][0]
['in', 'the', 'beginning', 'god', 'created', 'the', 'heaven', 'and', 'the', 'earth']
>>> bible[42][2][15]
['for', 'god', 'so', 'loved', 'the', 'world', 'that', 'he', 'gave', 'his', 'only', 'begotten', 'son']
```

```python
with open('kjv.txt', 'rt') as f:
    bible_d = list(f)

bible = []
...

for line in bible_d:
    if '#' in line:
        bible.append([line.replace('##','')])

import re

regex = re.compile('[^a-zA-Z ]')

chap = []
book = 0
chapter = 1
for line in range(len(bible_d)):
    clean = regex.sub('', bible_d[line].lower())
    if line == len(bible_d)-1:
        chap.append([clean])
        bible[book].append(chap)
    # Append verse to the same chapter
    elif bible_d[line].split(':')[0].strip('[') == str(chapter):
        chap.append([clean])
    # If higher chapter, append the previous one to the book and start a new one  
    elif bible_d[line].split(':')[0].strip('[') == str(chapter+1):
        bible[book].append(chap)
        chap = []
        chap.append([clean])
        chapter += 1
    # Connect new line to previous verse
    elif bible_d[line][0] == '\t':
        chap[-1][0] = chap[-1][0] + " " + clean
    # Start a new book         
    elif '#' in bible_d[line][0] and line != 0:
        bible[book].append(chap)
        chap = []
        chapter = 1
        book += 1

for book in bible:
    del book[0]

for book in bible:
    for chapter in book:
        for verse in range(len(chapter)):
            chapter[verse] = chapter[verse][0].split(' ')

```

```python
grader.check("q3a")
```

Output:
```
q3a results: All test cases passed!
```

**3(b)** (3 pts) In order to encode quickly encode a message, we need to be able to efficiently map a given word to all of its location(s) in the bible. Create a second data structure `bible_inv` such that `bible_inv[word]` contains a list of all the locations where that word occurs in `bible`. Each location should be encoded as a 4-tuple of integers, such that the following identity holds:
```
for b, c, v, w in bible_inv[word]:
    assert bible[b][c][v][w] == word
```

```python
bible_inv = {}
locations = []
for book_i in range(len(bible)):
    for chapter_i in range(len(bible[book_i])):
        for verse_i in range(len(bible[book_i][chapter_i])):
            for word_i in range(len(bible[book_i][chapter_i][verse_i])):
                word = bible[book_i][chapter_i][verse_i][word_i]
                tup = (book_i,chapter_i,verse_i,word_i)
                if word in bible_inv:
                    bible_inv[word].append(tup)
                else:
                    bible_inv[word] = []
                    bible_inv[word].append(tup)
...
```

Output:
```
Ellipsis
```

```python
grader.check("q3b")
```

Output:
```
q3b results: All test cases passed!
```

**3(c)** (3 pts) Finally, create the functions `enc_bible(s, bible_inv)` and `dec_bible(s, bible)` which encode and decode the string `s` given the mappings `bible` and `bible_inv` above. The output of `enc_bible` should be a string of hyphen-separated 4-tuples:

```
>>> m = enc_bible("the eagle flies at dawn", bible_inv)
>>> m
'5-20-42-7 25-16-2-9 18-77-44-5 25-23-17-10 39-27-0-10'
>>> dec_bible(m, bible)
"the eagle flies at dawn"
```

You may assume that any plaintext phrase `s` consists of entirely of unpunctuated lowercase words that exist in the bible.

**Note**: the output of `enc_bible` is not necessarily unique.

```python
def enc_bible(s, bible_inv):
    import random
    def tuple_to_key(tup):
        return "-".join(map(str, tup))
    return " ".join(tuple_to_key(random.choice(bible_inv[w])) for w in s.split(" "))

def dec_bible(s, bible):
    ret = []
    for word in s.split(" "):
        b, c, v, w = map(int, word.split("-"))
        ret.append(bible[b][c][v][w])
    return " ".join(ret)
```

```python
grader.check("q3c")
```

Output:
```
q3c results: All test cases passed!
```

**3(d)** (3 pts) One weakness of a book cipher is that it is not possible to encrypt a message containing words that are not found in the book. For example, I cannot encrypt the message `"i heart sushi"` using the bible cipher. Work around this limitation by creating a function `closest_encryptable(s, bible_inv)` which returns a copy of the plaintext `s` where each word in `s` that is not found in the bible has been replaced by its "nearest match" that *is* in the bible. 

```
>>> closest_encryptable("i heart sushi", bible_inv)
"i heart susi"
```

To measure the closeness between two words, we will use the function `difflib.SequenceMatcher`:

```python
from difflib import SequenceMatcher
print([SequenceMatcher(a="sushi", b=w).ratio() for w in ("shush", "pizza")])  # high vs low score
```

Output:
```
[0.8, 0.2]
```

In the event of a tie score, pick the word that comes first in alphabetical order. You can again assume that `s` is composed of unpunctuated lower-case words separated by spaces, as in the example.

```python
def closest_encryptable(s, bible_inv):
    encryption = []
    msg = s.split(' ')
    for w in msg:
        if w in bible_inv:
            encryption.append(w)
        else:
            ratios = {}
            for word in bible_inv:
                ratio = SequenceMatcher(a=w, b=word).ratio()
                ratios[word] = ratio

            closest = max(ratios.values())
            max_values = [a for a, b in ratios.items() if b == closest]
            max_values.sort(reverse=True)
            encryption.append(max_values[0])
            
    encryption = " ".join(encryption)
    return encryption
    ...
```

```python
grader.check("q3d")
```

Output:
```
q3d results: All test cases passed!
```

<!-- BEGIN QUESTION -->

### Question 4: A peculiar file format (5 pts)

The file `q4.dat` contains string data which has been encoded as follows:
- The first byte represents an unsigned integer $k_1$.
- The next $k_1$ bytes of the file represent $k_1$ ASCII characters of the message.
- The next character (the $k_1+1$-st) contains an unsigned integer $k_2$.
- The next $k_2$ bytes contain additional `data`. 
- ... and so on.

Decode the message.

Hints:
- `file.read(b)` reads `b` bytes from an open file `file`.
- To convert a byte to an integer, use the function `int.from_bytes()`. 
- To convert bytes to an ASCII string, use `bytes.decode('ascii')`.
- By printing out the result, it will be obvious if you have correctly decoded the message or not.

```python
f = open('q4.dat', 'rb')

my_bytes = []
byte = f.read(1)
my_bytes.append(byte)
while byte:
    byte = f.read(1)
    my_bytes.append(byte)
f.close

msg = []
cnt = 0

for i in range(len(my_bytes)):
    k = int.from_bytes(my_bytes[cnt],byteorder='little')
    if i < cnt+k+1 and i != 0:
        msg.append(my_bytes[i])
    elif i != 0:
        cnt = cnt + k + 1
        
mes = b''.join(msg)

print(mes.decode('ascii'))
...
```

Output:
```
,----,                    ,----,                          ,----,.                       ,----, 
                 ,/   .`|                  ,/   .`|                        ,'   ,' |                     .'   .`| 
  .--.--.      ,`   .'  : ,---,          ,`   .'  : .--.--.              ,'   .'   |    ,----..       .'   .'   ; 
 /  /    '.  ;    ;     /'  .' \       ;    ;     //  /    '.          ,----.'    .'   /   /   \    ,---, '    .' 
|  :  /`. /.'___,/    ,'/  ;    '.   .'___,/    ,'|  :  /`. /          |    |   .'    /   .     :   |   :     ./  
;  |  |--` |    :     |:  :       \  |    :     | ;  |  |--`           :    :  |--,  .   /   ;.  \  ;   | .'  /   
|  :  ;_   ;    |.';  ;:  |   /\   \ ;    |.';  ; |  :  ;_             :    |  ;.' \.   ;   /  ` ;  `---' /  ;    
 \  \    `.`----'  |  ||  :  ' ;.   :`----'  |  |  \  \    `.          |    |      |;   |  ; \ ; |    /  ;  /     
  `----.   \   '   :  ;|  |  ;/  \   \   '   :  ;   `----.   \         `----'.'\   ;|   :  | ; | ' 
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question -1 (extra credit; +5)

The file `secret.txt` included with this problem set contains an encrypted message. Can you decode it?

**Rules**: 

- In order to get credit, you must document every step that you used to arrive at the solution. Submissions without a plausible explanation of how the solution was derived will receive no points, regardless of correctness.
- You may collaborate with others _on this question only_. Your solution should list the name(s) of the your collaborators. Each collaborator will receive $5/(\text{num. collaborators})$ points in the event of a correct answer.
- You may not discuss this problem with anyone except your collaborators.

_Type your answer here, replacing this text._

<!-- END QUESTION -->



---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Output:
```
q1a results: All test cases passed!

q1b results: All test cases passed!

q1c results: All test cases passed!

q1d results: All test cases passed!

q1e results: All test cases passed!

q2a results: All test cases passed!

q2b results: All test cases passed!

q3a results: All test cases passed!

q3b results: All test cases passed!

q3c results: All test cases passed!

q3d results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Upload this .zip file to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Output:
```
<IPython.core.display.HTML object>
```

Tests for q1a (points: 2)
```
>>> p = primes()
>>> plist = [2,3,5,7,11,13,17,19,23,29]
>>> for i in range(10):
...     assert next(p) == plist[i]
```


Tests for q1b (points: 3)
```
>>> r = ruler()
>>> ruler_numbers = [1, 2, 2, 3, 3, 4, 4, 4, 5, 5, 5, 6, 6, 6, 7, 7, 7, 7, 8, 8, 8, 8]
>>> for n in ruler_numbers:
...     assert next(r) == n
```


Tests for q1c (points: 3)
```
>>> s = look_say()
>>> look_say_first_ten = ['1',
...  '11',
...  '21',
...  '1211',
...  '111221',
...  '312211',
...  '13112221',
...  '1113213211',
...  '31131211131221',
...  '13211311123113112211']
>>> for i in range(10):
...     s_i = next(s)
...     assert isinstance(s_i, str)
...     assert s_i == look_say_first_ten[i]
```


Tests for q1d (points: 3)
```
>>> first_five_months = ['July 2022', 'December 2023', 'March 2024', 'August 2025', 'May 2026']
>>> sm = special_months()
>>> for m in first_five_months:
...     assert next(sm) == m
```


Tests for q1e (points: 3)
```
>>> first_22_palindromic = ['February 2, 2022',
...  'February 20, 2022',
...  'February 21, 2022',
...  'February 22, 2022',
...  'February 23, 2022',
...  'February 24, 2022',
...  'February 25, 2022',
...  'February 26, 2022',
...  'February 27, 2022',
...  'February 28, 2022',
...  'March 2, 2023',
...  'March 20, 2023',
...  'March 21, 2023',
...  'March 22, 2023',
...  'March 23, 2023',
...  'March 24, 2023',
...  'March 25, 2023',
...  'March 26, 2023',
...  'March 27, 2023',
...  'March 28, 2023',
...  'March 29, 2023',
...  'April 2, 2024']
>>> pd = palindromic_dates()
>>> for pdate in first_22_palindromic:
...     assert next(pd) == pdate
```
```
>>> # BEGIN HIDDEN
>>> palindromic_103_110 = ['July 31, 2037',
...  'August 3, 2038',
...  'August 30, 2038',
...  'August 31, 2038',
...  'September 3, 2039',
...  'September 30, 2039',
...  'January 4, 2041']
>>> pd = palindromic_dates()
>>> for i in range(110):
...     pd_i = next(pd)
...     if i >= 103:
...         assert pd_i == palindromic_103_110[i-103]
```


Tests for q2a (points: 4)
```
>>> public_tests = [('thequickbrownfoxjumpsoverthelazydog', 1, 'uifrvjdlcspxogpykvnqtpwfsuifmbazeph'), ('howdoesthiscodehandlewrapping', 18, 'zgovgwklzakugvwzsfvdwojshhafy'),('thekeycanbenegative', -3, 'qebhbvzxkybkbdxqfsb')]
>>> for msg, k, enc in public_tests:
...     assert enc_caesar(msg, k) == enc
```
```
>>> public_tests = [('thequickbrownfoxjumpsoverthelazydog', 1, 'uifrvjdlcspxogpykvnqtpwfsuifmbazeph'), ('howdoesthiscodehandlewrapping', 18, 'zgovgwklzakugvwzsfvdwojshhafy'),('thekeycanbenegative', -3, 'qebhbvzxkybkbdxqfsb')]
>>> for msg, k, enc in public_tests: 
...     assert dec_caesar(enc, k) == msg
```


Tests for q2b (points: 4)
```
>>> public_tests = [('hello', 'sun', 'zyydi'),('goodbye', 'moon', 'sccqnms'),('dog', 'cats', 'foz'),('ilovestatistics', 'bayes', 'jlmzwttyxattggk')]
>>> 
>>> for msg, key, enc in public_tests:
...     assert enc_vignere(msg, key) == enc
```
```
>>> public_tests = [('hello', 'sun', 'zyydi'),('goodbye', 'moon', 'sccqnms'),('dog', 'cats', 'foz'),('ilovestatistics', 'bayes', 'jlmzwttyxattggk')]
>>> 
>>> for msg, key, enc in public_tests:
...     assert dec_vignere(enc, key) == msg
```


Tests for q3a (points: 3)
```
>>> assert bible[0][0][0] == ['in', 'the', 'beginning', 'god', 'created', 'the', 'heaven', 'and', 'the', 'earth']
>>> assert bible[42][2][15] == ['for', 'god', 'so', 'loved', 'the', 'world', 'that', 'he', 'gave', 'his', 'only', 'begotten', 'son', 'that', 'whosoever', 'believeth', 'in', 'him', 'should', 'not', 'perish', 'but', 'have', 'everlasting', 'life']
```


Tests for q3b (points: 3)
```
>>> for word in ['dawn', 'light', 'heaven']:
...     for b, c, v, w in bible_inv[word]:
...         assert bible[b][c][v][w] == word
```


Tests for q3c (points: 3)
```
>>> import re
>>> test_strs = ("the eagle flies at dawn", "heaven and earth reach him")
>>> for test_str in test_strs:
...     b = enc_bible(test_str, bible_inv)
...     pattern = re.compile("([0-9]+-){3}[0-9]+")
...     assert len(pattern.findall(b)) == 5, "Output of enc_bible not correctly formatted."
...     assert dec_bible(b, bible) == test_str
```


Tests for q3d (points: 3)
```
>>> assert closest_encryptable("i heart sushi", bible_inv) == "i heart susi"
```
```
>>> assert closest_encryptable("tiktok is so lit", bible_inv) == 'traitor is so lift'
```
